---
title: 'Wprowadzenie do pakietu R'
description: ""
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

R to język programowania, ale jednocześnie 

`@instructions`
fghdfghdfghfg

`@hint`


`@pre_exercise_code`
```{r}
# zmienna liczbowa 
liczba <- 1

#zmienna tekstowa
tekst <- "abcs"

```

`@sample_code`
```{r}
dfghdfg
```

`@solution`
```{r}
dfghdfghdfgh
```

`@sct`
```{r}
hdfgdhjkyuk
```
