---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.

`@instructions`
fghdfghdfghfg

`@hint`


`@pre_exercise_code`
```{r}
dfghdfghdfg
```

`@sample_code`
```{r}
dfghdfg
```

`@solution`
```{r}
dfghdfghdfgh
```

`@sct`
```{r}
hdfgdhjkyuk
```
